# Einsendeaufgabe 1: Numerische Genauigkeit und Gleichungssysteme (100 Punkte)
In dieser Aufgabe sollen Sie ein wenig mehr Erfahrung mit NumPy und numerischen Methoden gewinnen.  
Zur Erinnerung empfehle ich an dieser Stelle, die Definition der [IEEE-Fließkommazahlen](https://de.wikipedia.org/wiki/IEEE_754) zu wiederholen.  

## Addition von Zahlen (20 Punkte) 

Gegeben sei ein Array *array*, dass 100 mal die Zahl $10^{-16}$ enthält und einmal (als ersten Eintrag) die Zahl $1$. 

```python
import numpy as np

array = np.concatenate(([1], np.full(100, 1e-16)))
print(array)
```

Output:
```
[1.e+00 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16 1.e-16
 1.e-16]
```

**Aufgabe:** Addieren Sie alle Werte in `array`:
- mit der Funktion `np.sum()`
- mit einer Schleife, die von **vorne nach hinten** über `array` iteriert
- mit einer Schleife, die von **hinten nach vorne** über `array` iteriert

_Points:_ 6

```python
sum_numpy = np.sum(array)
print(f"{sum_numpy=}")

sum_forward = 0
for num in array:
    sum_forward += num
print(f"{sum_forward=}")

sum_reverse = 0
for num in reversed(array):
    sum_reverse += num
print(f"{sum_reverse=}")
```

Output:
```
sum_numpy=1.0000000000000084
sum_forward=1.0
sum_reverse=1.00000000000001
```

<!-- BEGIN QUESTION -->

**Aufgabe:** Erklären Sie die Ergebnisse. Wie werden die Zahlen 1 und 1e-17 im Computer dargestellt? (ausführliche Erläuterung erforderlich!)  

_Points:_ 12

In den meisten moderenen Computer-Architekturen werden Fließkommazahlen (Float) im Sinne des IEEE 754 Standard repräsentiert.
Floats werden oBdA in 64 bits repräsentiert:
- Sign bit: Zahls ist negativ oder positiv
- Exponent: 11 Bits die die Größenordung der Zahl angeben
- Mantissa: 52 bits die wesentliche Ziffern der Zahl repräsentieren
Der Wert einer Fließkommazahl in diesem Format kann mit der folgenden Funktion ermittelt werden:

```python
def ieee_754_to_float(sign, exponent, mantissa):
    exponent_value = exponent - 1023
    mantissa_value = 1 + mantissa / (2**52)
    float_value = (-1)**sign * mantissa_value * (2**exponent_value)
    return float_value
```

$10^{-16}$ ist eine sehr kleine Zahl, nahe dem Präzessions-Limit von Fließkommazahlen.
Die exakte Darstellung der Zahl im binären Format mit 64 bits ist auch für einen modernen Computer 
sehr herausfordernd und kann zu Rundungsfehlern führen.

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Aufgabe:** Welches Verfahren verwendet `np.sum()` für die Addition?

_Points:_ 2

Rekursive Paarweise Summierung

<!-- END QUESTION -->

## Grenzwerte von Funktionen (20 Punkte)

Geben sei die Funktion $f(x)=\frac{e^x-1}{x}$

Es gilt: $lim_{x \to 0}f(x) = 1$

**Aufgabe:** Schreiben Sie eine Python Funktion `f(x)`, mit deren Hilfe Sie $f(x)$ für $x= 10^{-1}, 10^{-2}, \cdots , 10^{-15}$ mit NumPy berechnen und geben Sie das Ergebnis aus.

_Points:_ 8

<!-- grade_id: cell-d5ed715a9745e7d6, points: 0, grade: false, solution: true -->
```python
def f(x):
    return (np.exp(x) - 1) / x if x != 0 else 1.0


input_numbers = [10**-i for i in range(1, 16)]
result_limits = [f(x) for x in input_numbers]
print(f"{result_limits=}")
```

Output:
```
result_limits=[1.0517091807564771, 1.005016708416795, 1.0005001667083846, 1.000050001667141, 1.000005000006965, 1.0000004999621837, 1.0000000494336803, 0.999999993922529, 1.000000082740371, 1.000000082740371, 1.000000082740371, 1.000088900582341, 0.9992007221626409, 0.9992007221626409, 1.1102230246251565]
```

<!-- grade_id: cell-86fd557d439b5a41, points: 0, grade: false, solution: true -->
```python
input_numbers_2 = [10**-i for i in range(16, 21)]
result_limits_2 = [f(x) for x in input_numbers]
print(f"{result_limits_2=}")
```

Output:
```
result_limits_2=[1.0517091807564771, 1.005016708416795, 1.0005001667083846, 1.000050001667141, 1.000005000006965, 1.0000004999621837, 1.0000000494336803, 0.999999993922529, 1.000000082740371, 1.000000082740371, 1.000000082740371, 1.000088900582341, 0.9992007221626409, 0.9992007221626409, 1.1102230246251565]
```

<!-- BEGIN QUESTION -->

**Aufgabe:** Was fällt auf? Welche Erklärung haben Sie für das Ergebnis? (Ausführliche Erklärung erforderlich) 

_Points:_ 12

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Gauß-Verfahren und LU-Zerlegung (60 Punkte) 
Obwohl man, wie wir sehen werden, lineare Gleichungssysteme einfach mit SciPy lösen kann, wollen wir in dieser Aufgabe  einmal selbst eine sogenannte LU-Zerlegung bzw. das Gauß-Verfahren implementieren. Außerdem werden wir diesen Code auch später noch verwenden, wenn wir uns mit Performance-Optimierung und Parallelisierung beschäftigen.

Das Gauß-Verfahren ist ein wichtiges Verfahren zum Lösen von linearen Gleichungssystemen, wir gehen hier schrittweise mit der Implementierung vor.

**Aufgabe:** Implementieren Sie die Funktion `b_solve` zur Lösung einer Dreiecksgleichung der Form:

$\left[ {\begin{array}{cccc} 
a_{11} & a_{12} & \cdots & a_{1n}\\ 
0 & a_{22} & \cdots & a_{2n}\\ 
\vdots & \vdots & \ddots & \vdots \\ 
0 & 0 & \cdots & a_{nn}\\ 
\end{array} } \right] \left[{\begin{array}{c}  
x_1\\ 
x_2\\ 
\vdots\\ 
x_n\\ 
\end{array}}\right] =  
\left[{\begin{array}{c}  
b_1\\ 
b_2\\ 
\vdots\\ 
b_n\\ 
\end{array}}\right]$ 

Hier beginnt man mit der Löung der letzten Zeile, denn es gilt:

$x_n = b_n/a_{nn} $
und dann

$x_{n-1} =  (b_{n-1}-a_{(n-1),n}*x_{n})/a_{(n-1),(n-1)}$

$x_{n-2} =  (b_{n-2}-(a_{(n-2),n}*x_{n}+a_{(n-2),(n-1)}*x_{n-1})/a_{(n-2),(n-2)}$


*Hinweis*: Achten Sie darauf, dass ihr Code nicht nur korrekt, sondern auch effizient ist. 

_Points:_ 12

<!-- grade_id: cell-045bae77f313bff4, points: 0, grade: false, solution: true -->
```python
import numpy as np


def b_solve(A, b):
    """Solve a system of linear equations A * x = b where A is an upper triangular matrix."""
    # Setup
    n = len(b)
    x = np.zeros(n)
    
    # Iterative calculations
    for i in range(n-1, -1, -1):
        if A[i, i] == 0:
            raise ValueError("Matrix is singular.")
        sum_ax = 0
        for j in range(i+1, n):
            sum_ax += A[i, j] * x[j]
        x[i] = (b[i] - sum_ax) / A[i, i]
    
    return x
```

**Aufgabe:** Implementieren Sie die LU-Zerlegung einer Matrix mit Pivotisierung.  
Die Funktion `lu_decomposition` soll die Matrizen L und U berechnen.  
Zu Beginn wird der Funktion eine Matrix A übergeben.  


Implementieren Sie die Funktion "in-place", d.h. die Ergebnisse für L und U werden direkt in A geschrieben.  
Dabei wird die ursprüngliche Matrix überschrieben.  
Die Funktion gibt den Pivot-Vektor zurück, der die folgende Form haben sollte:

$ 
P=\left[{\begin{array}{ccc} 
0 & 1& 0\\ 
0 & 0& 1\\ 
1 & 0& 0\\ 
\end{array} } \right] \Rightarrow \left[{\begin{array}{c}  
1\\ 
2\\ 
0\\ 
\end{array}}\right] 
$ 

D.h. es wird nur der Index des Wertes in jeder Zeile gespeichert, der nicht Null ist.

Wenn man dies macht, kann man später für die Lösung einfach den Vektor `bc = b[P]` verwenden, um auch $b$ richtig zu pivotisieren.  

Hinweise: 
- Die Pivot-Zeile lässt sich am einfachsten mit `np.argmax` finden. Die Zeile mit dem höchsten Betrag können Sie mit `np.abs` finden.

- Tauschen zweier Elemente in einem NumPy Array: `v[[a, b]] = v[[b, a]]`

- Versuchen Sie, die interne Vektorisierung von NumPy im Eliminationsschritt zu nutzen. Verwenden Sie dazu die Slicing-Syntax, um die innerste Schleife zu ersetzen.

_Points:_ 15

<!-- grade_id: cell-b482047b4df4ac5e, points: 0, grade: false, solution: true -->
```python
def lu_decomposition(A):
    """
    Implement LU-Decomposition of a square matrix.
    This function works in-place i.e. it stores information about L and U in A and returns the pivot vector.
    """
    n = A.shape[0]
    p = np.arange(n)  # Natural ordering

    for i in range(n):
        # Select the pivot index with the max. absolute value in column i
        max_row = i + np.argmax(np.abs(A[i:n, i]))
        
        # Swap the current row with the pivot row in matrix A and pivot vecor p
        if i != max_row:
            A[[i, max_row], :] = A[[max_row, i], :]
            p[[i, max_row]] = p[[max_row, i]]

        # Perform the elimination below the diagnal
        for j in range(i + 1, n):
            A[j, i] /= A[i, i]  # Modify the factor (L part)
            A[j, i+1:] -= A[j, i] * A[i, i+1:]  # Update remaining submatrix (U part)

    return p
```

<!-- BEGIN QUESTION -->

**Aufgabe:** Schreiben Sie eine Funktion `custom_solve`, welche mit Hilfe von `lu_decomposition` ein lineares Gleichungssystem der Form $A*x=b$ löst.  

Hinweis: Versuchen Sie, die interne Vektorisierung von NumPy zu nutzen. Verwenden Sie dazu die Slicing-Syntax, um die innerste Schleife zu ersetzen.

_Points:_ 12

<!-- grade_id: cell-e08e49ee4f4639ff, points: 0, grade: false, solution: true -->
```python
def custom_solve(A, b):
    """Solve a system of linear equations A * x = b and return result vector x."""
    A = A.copy() # Use a copy of A to keep original Data
    # Perform LU
    p = lu_decomposition(A)
    
    # Pivot b using the pivot vector
    b = b[p]
    
    n = A.shape[0]
    x = np.zeros_like(b)
    y = np.zeros_like(b)

    # Fwd substitution to solve L * y = b
    for i in range(n):
        y[i] = b[i] - np.dot(A[i, :i], y[:i])

    # Back substitution to solve U * x = y
    for i in range(n-1, -1, -1):
        x[i] = (y[i] - np.dot(A[i, i+1:], x[i+1:])) / A[i, i]

    return x
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Aufgabe:** Schreiben Sie eine Testfunktion, die überprüft, ob die Lösung ihrer Implementierung des Gauß-Verfahrens richtig ist.  
Testen Sie Gleichungssysteme der Größe $16 \times 16$, $32\times 32$ und $64 \times 64$.

Sie können zum Vergleich die SciPy Funktion `linalg.solve` verwenden:

_Points:_ 6

```python
# Beispiel linalg
import scipy.linalg as linalg

n = 128
rng = np.random.default_rng(seed=123)
A = rng.random((n, n))
b = rng.random(n)
c = linalg.solve(A, b)
```

<!-- grade_id: cell-f6a6bc3ba063c4a4, points: 0, grade: false, solution: true -->
```python
import scipy.linalg as linalg

# test different dimensions
def test_solve():
    sizes = [16, 32, 64]
    rng = np.random.default_rng(seed=123)
    
    for size in sizes:
        A = rng.random((size, size))
        b = rng.random(size)

        # custom solver
        x_custom = custom_solve(A, b)
        # scipy.linalg.solve
        x_scipy = linalg.solve(A, b)

        # Norm of the difference of the solutions
        diff = np.linalg.norm(x_custom - x_scipy)
        print(f"Size: {size}x{size} - Norm of difference: {diff}")
```

```python
test_solve()
```

Output:
```
Size: 16x16 - Norm of difference: 1.4045319291999208e-14
Size: 32x32 - Norm of difference: 3.8570912044911315e-14
Size: 64x64 - Norm of difference: 2.3563871603039596e-13
```

<!-- END QUESTION -->

**Aufgabe:** Wir wollen nun weiter die *fertigen* SciPy Funktionen betrachten. Mit den Funktionen `lu` und `lu_factor` aus dem Packet `scipy.linalg `  können Sie eine LU-Dekomposition durchführen [[1]](https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.lu_factor.html#scipy.linalg.lu_factor) und [[2]](https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.lu.html).  
Beide Funktionen haben unter anderem den Parameter `overwrite_a`, der `False` oder `True` sein kann. 

Evaluieren Sie die Performance (Laufzeit) der LU-Zerlegung einer 4096x4096 Matrix, wenn Sie jeweils die Funktionen `lu_factor` und `lu` mit `overwrite_a = True` oder `overwrite_a = False` verwenden.

In der Liste `minimums` sollen die minimalen Laufzeiten in folgender Reihenfolge stehen:
1. `lu_factor` mit `overwrite_a = False`
2. `lu_factor` mit `overwrite_a = True`
3. `lu` mit `overwrite_a = False`
4. `lu` mit `overwrite_a = True`

Hinweis: Da die Matrix A bei der Verwendung von `overwrite_a = True` verändert wird, kann es zu einer Vereinfachung der Berechnung kommen, wenn die Funktion mehrfach nacheinander ausgeführt wird. Aus diesem Grund muss die `timeit` Funktion [repeat](https://docs.python.org/3/library/timeit.html#timeit.repeat) mit `repeat = 2` und `number = 1`  verwendet werden. Die Matrix wird im `setup` Parameter erstellt (aus Zufallswerten mit dem vorgegebenen "random number generator").

_Points:_ 6

<!-- grade_id: cell-5978db060fba8a0d, points: 0, grade: false, solution: true -->
```python
from scipy.linalg import lu_factor, lu
from timeit import repeat

n = 4096 # Matrix: 4096x4096
number = 1
repetitions = 2
# store all four minimal runtimes in this list
minimums = []
rng = np.random.default_rng(seed=123)
```

```python
# create a 4096x4096 matrix
A = rng.random((n,n))
A
```

Output:
```
array([[0.68235186, 0.05382102, 0.22035987, ..., 0.33685154, 0.42171926,
        0.68634679],
       [0.16037319, 0.39680471, 0.56063436, ..., 0.24727999, 0.78631072,
        0.13823181],
       [0.88854304, 0.47116631, 0.65595772, ..., 0.82424156, 0.81338266,
        0.96395207],
       ...,
       [0.59107149, 0.67920093, 0.81553295, ..., 0.6467052 , 0.9576259 ,
        0.8987389 ],
       [0.41556502, 0.57935074, 0.14038325, ..., 0.88117271, 0.85576105,
        0.31965922],
       [0.73999952, 0.7111939 , 0.83864583, ..., 0.46134872, 0.11746248,
        0.87723509]])
```

```python
# lu_factor mit overwrite_a = False
setup_code = "from scipy.linalg import lu_factor; import numpy as np; rng = np.random.default_rng(seed=123); A = rng.random((4096, 4096))"
test_code = "lu_factor(A, overwrite_a=False)"
min_time = min(repeat(stmt=test_code, setup=setup_code, repeat=repetitions, number=number))
minimums.append(min_time)
minimums
```

Output:
```
[0.8741143730003387]
```

```python
# lu_factor mit overwrite_a = True
test_code = "lu_factor(A, overwrite_a=True)"
min_time = min(repeat(stmt=test_code, setup=setup_code, repeat=repetitions, number=number))
minimums.append(min_time)
minimums
```

Output:
```
[0.8741143730003387, 0.6791798719932558]
```

```python
# lu mit overwrite_a = False
setup_code = "from scipy.linalg import lu; import numpy as np; rng = np.random.default_rng(seed=123); A = rng.random((4096, 4096))"
test_code = "lu(A, overwrite_a=False)"
min_time = min(repeat(stmt=test_code, setup=setup_code, repeat=repetitions, number=number))
minimums.append(min_time)
minimums
```

Output:
```
[0.8741143730003387, 0.6791798719932558, 1.1892444280019845]
```

```python
# lu mit overwrite_a = True
test_code = "lu(A, overwrite_a=True)"
min_time = min(repeat(stmt=test_code, setup=setup_code, repeat=repetitions, number=number))
minimums.append(min_time)
minimums
```

Output:
```
[0.8741143730003387,
 0.6791798719932558,
 1.1892444280019845,
 1.0395183880027616]
```

**Aufgabe:** Untersuchen Sie den Speicherverbrauch der 4 Varianten mit memray.

Um den Einfluss der anderen Berechnungen zu minimieren, empfehlen wir hier, die Berechnungen in eine neue Datei zu schreiben und mit `!` in einer Code-Zelle memray's Command-Line Interface zu verwenden. 

Am besten tracen Sie nur die LU Funktion mit einem Context Manager!
Verwenden Sie jeweils folgende Dateinamen als `file_name` Parameter in `memray.Tracker`:

- lu_inplace.bin
- lu_notinplace.bin
- lu_factor_inplace.bin
- lu_factor_notinplace.bin

Beantworten Sie bitte die folgenden Fragen:
* Welche Version verwendet am meisten Speicher? 
* Bei welcher Version ist der Speicherverbrauch am konstantesten? 

_Points:_ 9

```python
%%file lu_inplace.py
import memray
import numpy as np
from scipy.linalg import lu

n = 4096
rng = np.random.default_rng(seed=123)
A = rng.random((n,n))

# use context manager
with memray.Tracker("lu_inplace.bin"):
    lu(A, overwrite_a=True)
```

Output:
```
Overwriting lu_inplace.py
```

```python
%%file lu_notinplace.py
import memray
import numpy as np
from scipy.linalg import lu

n = 4096
rng = np.random.default_rng(seed=123)
A = rng.random((n,n))

with memray.Tracker("lu_notinplace.bin"):
    lu(A, overwrite_a=False)
```

Output:
```
Overwriting lu_notinplace.py
```

```python
%%file lu_factor_inplace.py
import memray
import numpy as np
from scipy.linalg import lu_factor

n = 4096
rng = np.random.default_rng(seed=123)
A = rng.random((n,n))

with memray.Tracker("lu_factor_inplace.bin"):
    lu_factor(A, overwrite_a=True)
```

Output:
```
Writing lu_factor_inplace.py
```

```python
%%file lu_factor_notinplace.py
import memray
import numpy as np
from scipy.linalg import lu_factor

n = 4096
rng = np.random.default_rng(seed=123)
A = rng.random((n,n))

with memray.Tracker("lu_factor_not_inplace.bin"):
    lu_factor(A, overwrite_a=False)
```

Output:
```
Writing lu_factor_notinplace.py
```

```python
!python lu_inplace.py
```

```python
!python lu_inplace.bin
```

Output:
```
python: can't open file '/home/jovyan/work/lu_inplace.bin': [Errno 2] No such file or directory
```

```python
!memray run lu_notinplace.py
```

```python
!memray run lu_factor_inplace.py
```

```python
!memray run lu_factor_notinplace.py
```

```python
# Optional
!memray flamegraph lu_inplace.bin
```

Output:
```
/usr/bin/sh: 1: memray: not found
```

```python
# Delete files when ready
!rm -rf lu_* memray-*
```

```python

```

Tests for addition_ieee (points: 12)


Tests for addition_npsum (points: 2)


Tests for addition_variants (points: 6)


Tests for limits_explanation (points: 12)


Tests for limits_fx (points: 8)
```
>>> print(result_limits)
[1.05170918 1.00501671 1.00050017 1.00005    1.000005   1.0000005
 1.00000005 0.99999999 1.00000008 1.00000008 1.00000008 1.0000889
 0.99920072 0.99920072 1.11022302]
```
```
>>> print(result_limits_2)
[0. 0. 0. 0. 0.]
```


Tests for linearsystem_bsolve (points: 12)


Tests for linearsystem_customsolve (points: 12)


Tests for linearsystem_customsolve_test (points: 6)


Tests for linearsystem_ludecomp (points: 15)


Tests for linearsystem_scipy_memory (points: 9)


Tests for linearsystem_scipy_time (points: 6)
